# CMSI 185 - Extra Credit Assignment 
### DUE: 3 November 2020 by 11:59PM (Pacific Time)

- CMSI-185 assignments will be generated and submitted as Jupyter Notebooks
- Read the instructions for each problem and complete the assignment cell 
- i.e., the cells that contain comments like **YOUR CODE HERE** or **YOUR ANSWER HERE**
- Extra credit notebooks will be worth 10 points, and will be applied to your overall homework score. 

---

![image](https://upload.wikimedia.org/wikipedia/commons/5/5a/Vector-triangle-inequality.svg)

### Problem 1 - Conditionals and Branching
The Triangle Inequality Theorem states that the sum of two side lengths of a triangle is always greater than the third side. Thus, if any side is greater than or equal to the sum of the other two sides, the three lengths cannot form a triangle.

Write a function named `is_triangle` that accepts three integers as arguments that represent the length of three sticks. The function should return `True` if the 3 integers can form a triangle, otherwise the function should return `False`.

<!-- grade_id: cell-b01d9d94ee1d6b96, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-da72c4884e75fefc, points: 5, grade: true, solution: false -->
```python
# Test cases
assert is_triangle(5,5,5) == True
assert is_triangle(10,5,5) == False
```

---

![image](https://developers.google.com/edu/python/images/hello.png)

### Problem 2 - String Manipulation
Write a function named `reverse_string` that takes a string input of alpha characters and returns the reverse of the string, with uppercase and lowercase letters switched. For example, given the input `Practice Makes perfect` the function should return `TCEFREP SEKAm ECITCARp`. 

Your function does not need to handle numeric characters since the input strings will only consist of letters and whitespace characters.

Be sure to `return` the new string value from your function, not just print it. 

<!-- grade_id: cell-55bd748e3ddfca23, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-13ed72d01b963108, points: 5, grade: true, solution: false -->
```python
# Test cases
assert reverse_string("StuDent") == "TNEdUTs"
assert reverse_string("Practice Makes perfect") == "TCEFREP SEKAm ECITCARp"
```

---